# Homework 1: Data Types, Functions and Conditionals (29 points)

**STATS 507, Fall 2025**

-------

**Name: YI WANG** 

**Email: wangyii@umich.edu** 

**Time spent on this homework: 6 hours** 

**I did not discuss this homework with anyone.**



## Problem 1: Severance textbook exercises (8 points) 
Complete the following exercises from the textbook which you can find [here](https://www.py4e.com/html3/) for free. 

Note that some questions will ask you to get input from the user via the ```input```
function. Since Jupyter has no command-line interface, you will instead implement the program as a function. The answer cells you find below will provide a template to help you get started.

- Ch 1, Exercises 3,5,6,8  
- Ch 2, Exercises 3,4  
- Ch 3, Exercises 1, 2  
- Ch 4, Exercise 6  
- Ch 6, Exercise 1  

---
### Chapter 1, Exercises (2 points) 

#### Exercise 3 
What is the difference between a compiler and an interpreter?

<!-- grade_id: ch1-ex3-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable called 'ans'. 
# For example: "ans = 'This is a sample answer.'"
# You can use markdown formatting. 
ans = """
1. **Interpreter**: Reads source code, parses it, and executes it line by line.
2. **Compiler**: Reads source code, parses it, and converts it into machine language before execution.

**IN DETAIL:**

- An **interpreter** reads the source code of the program as written by the programmer, parses the source code, and interprets the instructions on the fly. 
It is the nature of an interpreter to be able to have an interactive conversation as shown above.
Python is an interpreter and when we are running Python interactively, we can type a line of Python (a sentence) and Python processes it immediately and is ready for us to type another line of Python.
- A **compiler** needs to be handed the entire program in a file, and then it runs a process to translate the high-level source code into machine language and then the compiler puts the resulting machine language into a file for later execution.
"""


```

<!-- grade_id: ch1-ex3-test, points: 0.5, grade: true, solution: false -->
```python
# Make sure you answer is rendered as markdown below. 
from IPython.display import Markdown as md
assert type(ans) == str
md(ans)
```

Output:
```
<IPython.core.display.Markdown object>
```

#### Exercise 5
What is wrong with the following code:
```
>>> primt 'Hello world!'
File "<stdin>", line 1
  primt 'Hello world!'
                     ^
SyntaxError: invalid syntax
>>>
```

<!-- grade_id: ch1-ex5-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable called 'ans'. 
# You can use markdown formatting. 
ans = """
There are two errors in this code:
1. **Typo**: 'primt' should be 'print' - there's a missing 'n'
2. **Python version syntax**: In Python 3, print is a function and requires parentheses

The correct code should be: `print('Hello world!')`
"""
```

<!-- grade_id: ch1-ex5-test, points: 0.5, grade: true, solution: false -->
```python
# Make sure you answer is rendered as markdown below. 
from IPython.display import Markdown as md
assert type(ans) == str
md(ans)
```

Output:
```
<IPython.core.display.Markdown object>
```

#### Exercise 6
Where in the computer is a variable such as `x` stored after the following Python line finishes?

a) Central processing unit

b) Main Memory

c) Secondary Memory

d) Input Devices

e) Output Devices

<!-- grade_id: cell-b0f556c01b7e2276, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable 'ans'. 
# Should be one of ['a', 'b', 'c', 'd', 'e']
ans = 'b'  # Main Memory
```

<!-- grade_id: cell-3176d2d916422842, points: 0.5, grade: true, solution: false -->
```python
assert type(ans) == str
assert ans in ['a','b','c','d','e']
```

#### Exercise 8
Explain each of the following using an example of a human capability: (1) Central processing unit, (2) Main Memory, (3) Secondary Memory, (4) Input Device, and (5) Output Device. For example, “What is the human equivalent to a Central Processing Unit”?

<!-- grade_id: ch1-ex8-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable called 'ans'. 
# You can use markdown formatting. 
# YOUR CODE HERE
ans = """
Human equivalents to computer components:

1. **Central Processing Unit (CPU)**: The brain - processes information, makes decisions, and performs calculations
2. **Main Memory (RAM)**: Short-term memory - temporarily holds information you're currently working with
3. **Secondary Memory**: Long-term memory - stores information permanently for later retrieval
4. **Input Device**: Senses - receive information from the external environment
5. **Output Device**: Voice and all kinds of expressions - communicate processed information to others
"""

```

<!-- grade_id: ch1-ex8-test, points: 0.5, grade: true, solution: false -->
```python
# Make sure you answer is rendered as markdown below. 
from IPython.display import Markdown as md
assert type(ans) == str
md(ans)
```

Output:
```
<IPython.core.display.Markdown object>
```

### Chapter 2, Exercises (2 points) 

#### Exercise 3
 Write a function which computes gross pay given the **hours** worked and the **rate_per_hour**. 
 
 (Note: You do not need to use the input() function as in the textbook question.)

<!-- grade_id: cell-312caeaaf6abe1a9, points: 0, grade: false, solution: true -->
```python
def get_gross_pay(hours, rate_per_hour):
    # YOUR CODE HERE
    return hours * rate_per_hour
```

<!-- grade_id: cell-458cb3b5b6333645, points: 1, grade: true, solution: false -->
```python
assert get_gross_pay(3, 10) == 30
```

#### Exercise 4
Assume that we execute the following assignment statements:
```
width = 17
height = 12.0
```
For each of the following expressions, write the value of the expression and the type (of the value of the expression).

1. `width//2`

2. `width/2.0`

3. `height/3`

4. `1 + 2 * 5`

<!-- grade_id: cell-e69aa106b6a89086, points: 0, grade: false, solution: true -->
```python
# Enter you answer as python list of tuples (value, type) in a variable 'ans'.
# For example, "ans = [(1,int), (2.0, float)]" (not actual answer)
# where the order of the list corresponds to order above.
# YOUR CODE HERE
ans = [(8,int), (8.5, float), (4.0, float), (11, int)]
```

<!-- grade_id: cell-796c8c2828e551ce, points: 1, grade: true, solution: false -->
```python
assert ans[0] == (8,int)
```

---
### Chapter 3, Exercises (2 points) 

#### Exercise 1 
Rewrite your `get_gross_pay` function from ch2, exercise 3 to give the employee 1.5 times the hourly rate for hours worked above 40 hours.

<!-- grade_id: cell-8f610a982f866fac, points: 0, grade: false, solution: true -->
```python
def get_gross_pay(hours, rate_per_hour):
    # YOUR CODE HERE
    if hours <= 40:
        return hours * rate_per_hour
    else:
        regular_pay = 40 * rate_per_hour
        overtime_hours = hours - 40
        overtime_pay = overtime_hours * rate_per_hour * 1.5
        return regular_pay + overtime_pay
```

<!-- grade_id: cell-537d9d1c5dcd290f, points: 1, grade: true, solution: false -->
```python
assert get_gross_pay(0, 1000)==0
assert get_gross_pay(2, 10)==20
```

#### Exercise 2
Rewrite your `get_gross_pay` program from ch2, exercise 3 using `try` and `except` so that your program handles non-numeric input gracefully by printing an informative message and returning nothing.

```
>>> get_gross_pay(20, 'nine')
Error, please enter numeric input

>>> get_gross_pay('forty', 'nine')
Error, please enter numeric input
```

<!-- grade_id: cell-b57ce5b94cb6cf92, points: 0, grade: false, solution: true -->
```python
def get_gross_pay(hours, rate_per_hour):
    # YOUR CODE HERE
    try:
        hours = float(hours)
        rate_per_hour = float(rate_per_hour)
        if hours <= 40:
            return hours * rate_per_hour
        else:
            regular_pay = 40 * rate_per_hour
            overtime_hours = hours - 40
            overtime_pay = overtime_hours * rate_per_hour * 1.5
            return regular_pay + overtime_pay
    except:
        print('Error, please enter numeric input')
        return None
```

<!-- grade_id: cell-582c5cae10097fd1, points: 1, grade: true, solution: false -->
```python
try: 
    get_gross_pay('nine', 'five')
    get_gross_pay('nine', 5)
    get_gross_pay(5, 5)
    get_gross_pay(5, 'nine')
    assert True
except: 
    assert False
```

Output:
```
Error, please enter numeric input
Error, please enter numeric input
Error, please enter numeric input
```

### Chapter 4, Exercises (1 point)

#### Exercise 1 
What is the purpose of the `def` keyword in Python?

a) It is slang that means “the following code is really cool”

b) It indicates the start of a function

c) It indicates that the following indented section of code is to be stored for later

d) b and c are both true

e) None of the above

<!-- grade_id: ch4-ex1-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable 'ans'. 
# Should be one of ['a', 'b', 'c', 'd', 'e']
# YOUR CODE HERE
ans = 'd' 
```

<!-- grade_id: ch4-ex1-test, points: 1, grade: true, solution: false -->
```python
assert type(ans) == str
assert ans in ['a','b','c','d','e']
```

### Chapter 6, Exercises (1 point) 

#### Exercise 6
Read the documentation of the string methods at https://docs.python.org/library/stdtypes.html#string-methods You might want to experiment with some of them to make sure you understand how they work. `strip` and `replace` are particularly useful.

<!-- grade_id: cell-9a484821ec48c9fc, points: 0, grade: false, solution: true -->
```python
# Create a variable 'ans' with the string: 'I read the documentation.'
ans = 'I read the documentation.'
```

<!-- grade_id: cell-4dd3c9dad6369e20, points: 1, grade: true, solution: false -->
```python
assert ans == 'I read the documentation.'
```

## Problem 2: Euclid's Algorithm (3 points)

Euclid’s algorithm (https://en.wikipedia.org/wiki/Euclidean_algorithm) is a method
for finding the greatest common divisor (GCD) of two integers. Recall that the GCD of
two integers $m$ and $n$ is the largest number that divides both $m$ and $n$.

---
### Part 1 (2 points)

The __Wikipedia__ page above includes several pseudocode implementations of Euclid’s
algorithm. Choose one of these, and use it to implement a function `gcd`, which takes
two integers as its arguments and returns their GCD. You may assume that both
inputs are postive integers (i.e. $m,n \in \mathbb Z$ and $m, n > 0$), so there is no need to include any error checking in your function.

**Note:** this is one of the rare occasions where you have my explicit permission to
look up your answer. Unless otherwise stated (e.g., as in this problem), looking up
solutions on Wikipedia or in any other non-class resource will be considered cheating!

<!-- grade_id: gcd, points: 0, grade: false, solution: true -->
```python
def gcd(a,b):
    """Compute the great common divisor (GCD) of two integers a,b. 
    """
    # YOUR CODE HERE
    while b != 0:
        a, b = b, a % b
    return abs(a) 
```

<!-- grade_id: gcd_tests, points: 2, grade: true, solution: false -->
```python
"""Check that your gcd function returns the correct output for several inputs"""
assert gcd(4,5) == 1
assert gcd(20,10) == 10
```

### Part 2 (1 point) 
What does your function do if one or both of the arguments are negative?

<!-- grade_id: cell-9b7a530cbce57b8d, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable called 'ans'. 
# You can use markdown formatting. 
# YOUR CODE HERE
ans = 'Works correctly! When dealing with negative inputs, the remainder calculation operates normally, while the absolute value operation guarantees a positive result.'

```

<!-- grade_id: cell-2bdb84d0d76cdb07, points: 1, grade: true, solution: false -->
```python
# Make sure you answer is rendered as markdown below. 
from IPython.display import Markdown as md
assert type(ans) == str
md(ans)
```

Output:
```
<IPython.core.display.Markdown object>
```

## Problem 3: Approximating Euler's number $e$ (8 points) 

The base of the natural logarithm, $e$, is typically deﬁned as the inﬁnite sum below:
$$
\begin{align}
    e = \sum_{k=0}^{\infty} \frac{1}{k!} = 1 + 1 + \frac{1}{2} + \frac{1}{6} + \frac{1}{24} + \dots && (1) \\
\end{align}
$$

The term $k!$ denotes the factorial of $k$, which is
$$
k! = k \cdot (k-1) \cdot (k-2) ... 3 \cdot 2 \cdot 1 
$$

By convention, $0! = 1$ For more on Euler’s number, see [here](https://en.wikipedia.org/wiki/E_(mathematical_constant)). In this problem, we will explore diﬀerent approaches to approximating this number.

### Part 1 (2 points) 

An early characterization of Euler’s number, due to Jacob Bernoulli, was as the limit of

$$ 
(1 + \frac{1}{x})^{x}
$$

as $x \rightarrow \infty$. Deﬁne a function called `euler_limit` that takes as an argument an integer $n$, and returns a `ﬂoat` that approximates $e$ by taking $x = n$ in Equation (2). You may assume that the input to your function will be a positive integer (of course, your function will still run just ﬁne if `n` is, say, a positive ﬂoat, but we will only use integer values in what follows).

<!-- grade_id: euler_limit, points: 0, grade: false, solution: true -->
```python
def euler_limit(n): 
    """Takes in a positive integer n and returns a float approximating the natural number e."""
    # YOUR CODE HERE
    return (1 + 1/n) ** n
```

<!-- grade_id: test_euler_limit1, points: 1, grade: true, solution: false -->
```python
"""Check that your euler_limit function returns the correct output for several inputs"""
assert euler_limit(1) == 2.0 
assert euler_limit(2) == 2.25
```

<!-- grade_id: test_euler_limit2, points: 1, grade: true, solution: false -->
```python
assert round(euler_limit(3),3) == 2.370
```

### Part 2 (1 point) 
What happens when you call `euler_limit(n)` for really huge values of `n` (say, $10^{16}$ or $10^{18}$ )? Why does this happen? 

**Hint:** the answer has to do with ﬂoating point arithmetic. Think about how $1/n$ is represented on your computer when $n$ is big.

<!-- grade_id: prob3-part2-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable called 'ans'. 
# You can use markdown formatting. 
# YOUR CODE HERE
ans = """
For very large values of n (like 10^16 or 10^18), the function returns 1.0 instead of 
approaching e ≈ 2.718.
This happens due to **floating point precision limitations**:
1. When n is extremely large, 1/n becomes very small
2. Due to finite precision in floating point representation, 1/n may become so small 
   that 1 + 1/n is rounded to exactly 1.0
3. Then (1.0)^n = 1.0, regardless of how large n is
This demonstrates the importance of numerical stability in mathematical computations.
"""
```

<!-- grade_id: prob3-part2-tests, points: 1, grade: true, solution: false -->
```python
# Make sure you answer is rendered as markdown below. 
from IPython.display import Markdown as md
assert type(ans) == str
md(ans)
```

Output:
```
<IPython.core.display.Markdown object>
```

### Part 3 (2 points) 
Deﬁne a function called `euler_sum` that takes a single non-negative integer argument `n`, and returns an approximation to $e$ based on the ﬁrst `n` terms of the sum in Equation (1). Your function should return a ﬂoat. 
You may assume that the input will be a non-negative integer, so you do not need to include error checking in your function. 

As an example, `euler_sum(4)` should return the sum of the ﬁrst four terms in Equation 1, so that `euler_sum(4)` returns $1 + 1 + 1/2 + 1/6 \approx 2.667$. 

**Note:** the sum in Equation 1 starts counting with $k = 0$ (i.e., it is “0-indexed”), while our function starts counting with $n = 1$ (i.e., it is “1-indexed”). `euler_sum(1)` should use one term from Equation (1), so that `euler_sum(1)` returns `1`. Similarly, `euler_sum(0)` should return `0`, since by convention an empty sum is equal to zero. 

**Note:** you may use the math.factorial function to compute $k!$, but I recommend, for the sake of practice, implementing the factorial function yourself, as it is a nice example of a problem that is easily solved with recursion.

<!-- grade_id: euler-sum, points: 0, grade: false, solution: true -->
```python
def euler_sum(n):
    def factorial(k):
        if k == 0 or k == 1:
            return 1
        result = 1
        for i in range(2, k+1):
            result *= i
        return result
    
    s = 0
    for k in range(n):  # k从0到n-1
        s += 1/factorial(k)
    return s
```

<!-- grade_id: test-euler-sum1, points: 1, grade: true, solution: false -->
```python
assert euler_sum(0) == 0 
assert round(euler_sum(4),3) == 2.667
```

<!-- grade_id: test-euler-sum2, points: 1, grade: true, solution: false -->
```python
assert euler_sum(3) == 2.5
```

### Part 4 (1 point) 
Deﬁne a function called `euler_approx` that takes a single argument, a positive ﬂoat `epsilon`, and uses the sum in Equation (1) to obtain an approximation of $e$ that is within `epsilon` of the true value of $e$. 

**Hint:** use a while-loop. 

**Note:** you can use the Python math module to get the true value of $e$ (up to ﬂoating point accuracy): `math.exp(1)`.

<!-- grade_id: euler-approx, points: 0, grade: false, solution: true -->
```python
import math

def euler_approx(epsilon):
    true_e = math.exp(1)
    n = 1
    
    while True:
        approx = euler_sum(n)
        if abs(approx - true_e) <= epsilon:
            return approx
        n += 1
```

<!-- grade_id: euler-approx-tests, points: 1, grade: true, solution: false -->
```python
assert abs(euler_approx(1e0) - math.exp(1)) < 1e0
```

### Part 5 (1 point) 
For a value of interest $x$ and an approximation $\hat x$, we say the difference $\hat x - x$ is the **error** of $f(k)$.

Deﬁne functions called `print_euler_sum_error` and `print_euler_lim_error` that take a single positive integer `n` as input. For $k=1,\dots,n$, the function will print the error of the approximation of $e$ given by `euler_sum` or `euler_limit`, respectively, evaluated at $k$. Each error should be printed on its own line (for $n$ lines in total).

<!-- grade_id: prob3-ex5-func, points: 0, grade: false, solution: true -->
```python
def print_euler_sum_error(n):
    true_e = math.exp(1)
    for i in range(1, n+1):
        approx = euler_sum(i)
        error = abs(approx - true_e)
        print(f"n={i}: euler_sum={approx:.6f}, error={error:.6f}")

def print_euler_lim_error(n):
    true_e = math.exp(1)
    for i in range(1, n+1):
        approx = euler_limit(i)
        error = abs(approx - true_e)
        print(f"n={i}: euler_limit={approx:.6f}, error={error:.6f}")
```

<!-- grade_id: cell-fccdb9f21de162ca, points: 1, grade: true, solution: false -->
```python
# Run this cell to check your functions. 
print_euler_sum_error(10)
print()
print_euler_lim_error(10)
```

Output:
```
n=1: euler_sum=1.000000, error=1.718282
n=2: euler_sum=2.000000, error=0.718282
n=3: euler_sum=2.500000, error=0.218282
n=4: euler_sum=2.666667, error=0.051615
n=5: euler_sum=2.708333, error=0.009948
n=6: euler_sum=2.716667, error=0.001615
n=7: euler_sum=2.718056, error=0.000226
n=8: euler_sum=2.718254, error=0.000028
n=9: euler_sum=2.718279, error=0.000003
n=10: euler_sum=2.718282, error=0.000000

n=1: euler_limit=2.000000, error=0.718282
n=2: euler_limit=2.250000, error=0.468282
n=3: euler_limit=2.370370, error=0.347911
n=4: euler_limit=2.441406, error=0.276876
n=5: euler_limit=2.488320, error=0.229962
n=6: euler_limit=2.521626, error=0.196655
n=7: euler_limit=2.546500, error=0.171782
n=8: euler_limit=2.565785, error=0.152497
n=9: euler_limit=2.581175, error=0.137107
n=10: euler_limit=2.593742, error=0.124539
```

### Part 6 (1 point) 
Which approximation is better? 

<!-- grade_id: prob3-part6-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable called 'ans'. 
# You can use markdown formatting. 
# YOUR CODE HERE
ans = """
The **euler_sum** 

Reasons:
1. **Faster convergence**: euler_sum reaches high accuracy with fewer terms
2. **Better precision**: For the same n, euler_sum typically has much smaller error
3. **Numerical stability**: euler_sum doesn't suffer from the floating point issues that affect euler_limit for large n

"""
```

<!-- grade_id: prob3-part6-test, points: 1, grade: true, solution: false -->
```python
# Make sure you answer is rendered as markdown below. 
from IPython.display import Markdown as md
assert type(ans) == str
md(ans)
```

Output:
```
<IPython.core.display.Markdown object>
```

## Problem 4: Testing properties of integers (5 points)
In this problem, you’ll get a bit more practice working with conditionals, and a ﬁrst exposure to the kind of thinking that is required in a typical “coding interview” question. A positive integer $n$ is a power of 2 if $n = 2^p$ for some integer $p \geq 0$.

### Part 1 (2 points) 
Write a function `is_power_of_2` that takes a positive integer as its only argument and returns a Boolean indicating whether or not the input is a power of 2. 

You may assume that the input is a positive integer. You may not use the built-in `math.sqrt` function in your solution. You should need only the division and modulus (%) operations. 

**Hint:** the simplest solution to this problem makes use of recursion, though recursion is not necessary.

<!-- grade_id: is_power_of_2, points: 0, grade: false, solution: true -->
```python
def is_power_of_2(n):
    if n == 1:
        return True
    if n < 1 or n % 2 != 0:
        return False
    return is_power_of_2(n // 2)
```

<!-- grade_id: is_power_of_2_test1, points: 1, grade: true, solution: false -->
```python
assert is_power_of_2(1)
assert is_power_of_2(2)
assert not is_power_of_2(15)
```

<!-- grade_id: is_power_of_2_test2, points: 1, grade: true, solution: false -->
```python
assert is_power_of_2(4)
```

### Part 2 (3 points)
Generalize your previous solution to a function `is_power` that takes two positive integers as its arguments, `b` and `n`, in that order, and returns a Boolean. `is_power(b,n)` should return `True` if `n` is a power of `b` and `False` otherwise.

<!-- grade_id: is_power, points: 0, grade: false, solution: true -->
```python
def is_power(b,n):
    # YOUR CODE HERE
    if b == 1:
        if n == 1:
            return True
        else:
            return False
    else:    
        if n == 1:
            return True
        if n < 1 or n % b != 0:
            return False
        return is_power(b,n//b)

```

<!-- grade_id: is_power_test1, points: 1.5, grade: true, solution: false -->
```python
assert is_power(5, 125)
assert not is_power(5, 150)
```

<!-- grade_id: is_power_test2, points: 1.5, grade: true, solution: false -->
```python
assert not is_power(1, 2)
assert is_power(3, 1)
```

## Problem 5: Create and share your stats 507 gitHub repository (5 points)

Create a new GitHub repository for the Stats 507 course and complete the following tasks:
### Set up the repository
Name the repository appropriately (e.g., "stats507-coursework");
Ensure the repository is public (unless otherwise specified)

### Add a README file
Create a file named "README.md" in the root directory;
Include a brief description of the course and the purpose of the repository

### Provide the repository link
Note: You should only post HW after the due date (or make it private)

## REPO LINK: [STATS-507-Fishwork](https://github.com/WangYii999/STATS-507-Fishwork)

<!-- grade_id: cell-9a484821ec48c9fc, points: 0, grade: false, solution: true -->
```python
# YOUR REPO LINK HERE
github_repo_link = "https://github.com/WangYii999/STATS-507-Fishwork"
## markdown link is in the previous cell
```